# Python Activity
## Introduction to NumPy

This notebook is designed to acquaint you with the NumPy module in Python. Refer to the content in Chapter 4 of _**Python for Data Analysis (3rd Ed.)**_ for examples of the type of code you need for these exercises.

For EACH exercise:

1. Read the description of the task
2. Type your solution in the code cell marked ```### YOUR CODE HERE```
3. Run your code (fix any issues and re-run if needed)
4. Run the TEST CELL that FOLLOWS your code cell. **_DO NOT MODIFY THE TEST CELL._**

The output from the TEST CELL will indicate whether you have performed the task correctly. If the result does not say _`Passed!`_ then you should return to your code cell and revise your code.

### What is NumPy?

[Numpy](http://www.numpy.org/) is a Python library that provides fast, efficient functions and commands for working with multidimensional arrays. NumPy is well-suited to implementing numerical linear algebra algorithms; for storing and operating on strictly numerical data, it is often much faster than Python's native data types (such as list and dictionary).

Some of the material in this activity is adapted from the more extensive tutorial found [here](http://www.scipy-lectures.org/intro/numpy/index.html)

### Load NumPy and Confirm Version


```python
import numpy as np
print(np.__version__)
```

##### Naming the Library 

By "naming" NumPy **`as np`** (see above) we can identify and use anything in the NumPy library with the prefix `np.` as shown in the examples below.

### Creating a NumPy Array

We start by creating a simple array using NumPy's `array` data type. This is a _one-dimensional_ array.

```python
a = np.array([2,4,6,8,10])
print(a)
```

### Creating a Two-Dimensional Array

Numpy supports multidimensional arrays. To have more than one dimension, nest each new dimension within a list when creating the array. An example is shown below.

```python
# Create a two-dimensional array with 3 rows and 4 columns (called a 3 x 4 array)
B = np.array([[0, 1, 2, 3],
              [4, 5, 6, 7],
              [8, 9, 10, 11]])

print(B)
```

#### Properties of an Array

A NumPy array has several _properties_ including `ndim`, `shape`, and `len`. 

Try the code cell below and see if you can determine the meaning of each property.

```python
print(B.ndim)    # What does this do?
print(B.shape)   # What does this do?
print(len (B))   # What does this do?
```

### NumPy Functions to Create Specific Matrices

Certain well-defined matrices can be created easily in Python (as 2-dimensional arrays). See examples below.

```python
# All elements set to 0
print(np.zeros((3, 4)))
```

```python
# All elements set to 1
print(np.ones((4, 3)))
```

```python
# The identity matrix, I
print(np.eye(4))
```

```python
# A diagonal matrix
print(np.diag([3, 4, 5]))
```

### Creating an Array with More Dimensions

The next example shows a nice way to create an array with more than 2 dimensions. It is a useful method because these higher dimension arrays are more difficult to visualize. Pay close attention to the `ndim`, `shape`, and `len` properties to be sure you understand what is going on.

```python
C1 = [[0, 1, 2, 3],
      [4, 5, 6, 7],
      [8, 9, 10, 11]]

C2 = [[12, 13, 14, 15],
      [16, 17, 18, 19],
      [20, 21, 22, 23]]

C = np.array([C1, C2])

print(C)
print(C.ndim)
print(C.shape)
print(len (C))
```

### Indexing and Slicing

Recall that **_index values start at zero_**, so an array with 3 elements will have index values 0, 1, and 2.

Study the examples below for a better understanding of how indexing and slicing work in Python. These examples use the 3-dimensional array $C$ that was created above.

> **Suggestion:** BEFORE running each cell, try to predict the output. Then run the cell. If your prediction was correct, great! If not, spend some time making sure you understand how the indexing and slicing produced the output that was given.

##### Index/Slice Example 1

```python
print (C[1,2,0])
```

##### Index/Slice Example 2

```python
print (C[0, 2, :])
```

##### Index/Slice Example 3

```python
print (C[1, :, 2])
```

##### Index/Slice Example 4

```python
print (C[:, :, 3])
```

##### Index/Slice Example 5

```python
print (C[1, -1, -2])
```

##### Index/Slice Example 6

```python
print (C[1, 0, ::-1])
```

##### Index/Slice Example 7

```python
print (C[1, 1, 1::2])
```

### Exercises

**Exercise 1.** Consider the following $6 \times 6$ matrix, which has 4 different subsets highlighted.

<img src="python_slicing_matrix.png">

For each subset illustrated above, write an indexing or slicing expression that extracts the subset. Store the result of each slice into `Z_green`, `Z_red`, `Z_gold`, and `Z_purple`.

<!-- grade_id: x_float, points: 0, grade: false, solution: true -->
```python
Z= np.array([[0,1,2,3,4,5],[10,11,12,13,14,15],[20,21,22,23,24,25],[30,31,32,33,34,35],[40,41,42,43,44,45],[50,51,52,53,54,55]])

# Construct `Z_green`, `Z_red`, `Z_gold`, and `Z_purple`:
###
### YOUR CODE HERE
###
```

<!-- grade_id: x_float_test, points: 1, grade: true, solution: false -->
```python
# Test cell: `check_Z`

print("==> Z:\n", Z)
assert (Z == np.array([np.arange(0, 6),
                       np.arange(10, 16),
                       np.arange(20, 26),
                       np.arange(30, 36),
                       np.arange(40, 46),
                       np.arange(50, 56)])).all()

print("\n==> Gold slice:\n", Z_gold)
assert (Z_gold == np.array ([3, 4])).all()

print("\n==> Red slice:\n", Z_red)
assert (Z_red == np.array ([2, 12, 22, 32, 42, 52])).all()

print("\n==> Purple slice:\n", Z_purple)
assert (Z_purple == np.array ([[44, 45], [54, 55]])).all()

print("\n==> Green slice:\n", Z_green)
assert (Z_green == np.array ([[20, 22, 24], [40, 42, 44]])).all()

print("\n(Passed!)")
```

#### Common Functions

Several functions can be applied to an array. A few are shown below. Notice that some produce another array of the same size, whereas others are _**aggregate**_ functions.


```python
D = np.array([[4, -1, 12, 6],
              [-5, 7, 2, 0],
              [3, 9, 8, -15]])
print(D)
```

##### Absolute Value

```python
np.abs(D)
```

##### Trig Functions
The sine function is just one of the trig functions available.

```python
np.sin(D)
```

##### Log and Exponential Functions
* The `exp` function returns $e^x$.
* The `log` function on our matrix $D$ would give some undefined results for 0 and negative values.

```python
np.exp(D)
```

#### Aggregate Functions

```python
np.mean(D)
```

```python
np.sum(D)
```

```python
np.max(D)
```

```python
np.min(D)
```

#### Using Aggregate Functions on a Specific Dimension ('Axis')

Recall the shape of our $D$ matrix:

```python
print(D.shape)
```

##### How the Shape Relates to Each Axis
* The first value (3) tells how many rows are in the matrix
* The second value (4) tells how many columns are in matrix
* The first axis (indexed at 0) moves across the rows
* The second axis (indexed at 1) moves across the columns

For instance, in the first example below, a mean is given for each _**column**_ because we use the values _**in each row**_ of that column. The first column of $D$ has values 4, -5, and 3 (on rows 1, 2, and 3, respectively).  The mean of those values is $\frac{2}{3} \approx 0.667$.

Similarly, in the second example, we are taking the mean of values in each column of a given row; the 4 values in the first row have a mean of 5.25.

```python
np.mean(D,axis=0)
```

```python
np.mean(D,axis=1)
```

**Exercise 2.**  Using matrix $D$ defined above, create an array of 3 values by computing the sum of all values in each column. Store the result into a variable called `my_sums`. Then compute the mean of the 3 sums, storing the result into `mean_sum`. 

<!-- grade_id: cell-a117d6495f42b850, points: 0, grade: false, solution: true -->
```python
#
# YOUR CODE HERE
#

```

<!-- grade_id: report_exam_avg_test, points: 1, grade: true, solution: false -->
```python
# Test cell: mean sums

def is_number(x):
    from numbers import Number
    return isinstance(x, Number)

assert my_sums.shape == (3,)
assert np.sum(my_sums)==30

assert is_number(mean_sum) 
assert mean_sum == 10

print("\n(Passed!)")
```

#### READY TO SUBMIT?
You've reached the end of this notebook. Be sure to restart and run all cells again to **make sure all cells are working** when they run in order. Then submit your **completed** HTML to the submission  folder for this activity.