# Python Introduction Activity
## Variables, data types, functions, and data tables

This notebook is designed to acquaint you with some fundamentals of Python. Refer to the content in Chapters 2 and 3 of _**Python for Data Analysis (3rd Ed.)**_ for examples of the type of code you need for these exercises.

For EACH exercise:

1. Read the description of the task
2. Type your solution in the code cell marked ```### YOUR CODE HERE```
3. Run your code (fix any issues and re-run if needed)
4. Run the TEST CELL that FOLLOWS your code cell. **_DO NOT MODIFY THE TEST CELL._**

The output from the TEST CELL will indicate whether you have performed the task correctly. If the result does not say _`Passed!`_ then you should return to your code cell and revise your code.

> **Hint: Read the test code!** The test code cell checks the output of your code to see whether it appears to produce correct results. You can often learn a lot by reading the test code. In fact, sometimes it gives you a hint about how to approach the problem. As such, you are encouraged to try to read the test cells; you can learn what result is expected, AND learn some additional tricks in Python.

**Exercise 1** (`x_float_test`). Create a variable named `x_float` whose numerical value is one (1) and whose type is *floating-point*.

<!-- grade_id: x_float, points: 0, grade: false, solution: true -->
```python
#
# YOUR CODE HERE
#

```

<!-- grade_id: x_float_test, points: 1, grade: true, solution: false -->
```python
# `x_float_test`: Test cell
assert x_float == 1
assert type(x_float) is float
print("\n(Passed!)")
```

**Exercise 2** (`report_exam_avg_test`). Let `a`, `b`, and `c` represent three exam scores as numerical values. Complete the function, `report_exam_avg(a, b, c)` so that it computes the average score (equally weighted) and returns the string, `'Your average score: XX'`, where `XX` is the average rounded to one decimal place. For example:

```python
    report_exam_avg(100, 95, 80) == 'Your average score: 91.7'
```

#### Note:
**_The function_ ```is_number``` _is provided for you in your code cell.  You should NOT need to modify it.  Simply fill in the code for the the_ ```report_exam_avg``` _function._**

<!-- grade_id: cell-a117d6495f42b850, points: 0, grade: false, solution: true -->
```python
def is_number(x):
    """Returns `True` if `x` is a number-like type, e.g., `int`, `float`, `Decimal()`, ..."""
    from numbers import Number
    return isinstance(x, Number)

def report_exam_avg(a, b, c):
    assert is_number(a) and is_number(b) and is_number(c)
    #
    # YOUR CODE HERE
    #

```

<!-- grade_id: report_exam_avg_test, points: 1, grade: true, solution: false -->
```python
# `report_exam_avg_test`: Test cell
msg = report_exam_avg(100, 95, 80)
print(msg)
assert msg == 'Your average score: 91.7'

print("Checking some additional randomly generated cases:")
for _ in range(10):
    ex1 = random() * 100
    ex2 = random() * 100
    ex3 = random() * 100
    msg = report_exam_avg(ex1, ex2, ex3)
    ex_rounded_avg = float(msg.split()[-1])
    abs_err = abs(ex_rounded_avg*3 - (ex1 + ex2 + ex3)) / 3
    print("{}, {}, {} -> '{}' [{}]".format(ex1, ex2, ex3, msg, abs_err))
    assert abs_err <= 0.05

print("\n(Passed!)")
```

**Exercise 3** (`floor_fraction_test`). Suppose you are given two variables, `a` and `b`, whose values are the real numbers, $a \geq 0$ (non-negative) and $b > 0$ (positive). Complete the function, `floor_fraction(a, b)` so that it returns $\left\lfloor\frac{a}{b}\right\rfloor$, that is, the *floor* of $\frac{a}{b}$. The *type* of the returned value must be `int` (an integer).

#### NOTE
**_This function uses the_ ```is_number``` _function defined in Exercise 2. Therefore, you need to successfully run the Exercise 2 code before attempting this exercise._**

<!-- grade_id: floor_fraction, points: 0, grade: false, solution: true -->
```python
def floor_fraction(a, b):
    assert is_number(a) and a >= 0
    assert is_number(b) and b > 0
    #
    # YOUR CODE HERE
    #

```

<!-- grade_id: floor_fraction_test, points: 1, grade: true, solution: false -->
```python
# `floor_fraction_test`: Test cell
from random import random
a = random()
b = random()
c = floor_fraction(a, b)

print('floor_fraction({}, {}) == floor({}) == {}'.format(a, b, a/b, c))
assert b*c <= a <= b*(c+1)
assert type(c) is int
print('\n(Passed!)')
```

**Exercise 4** (`minmax_test`). Complete the function `minmax(L)`, which takes a list `L` and returns a pair---that is, 2-element Python tuple, or "2-tuple"---whose first element is the minimum value in the list and whose second element is the maximum. For instance:

```python
  minmax([8, 7, 2, 5, 1]) == (1, 8)
```

```python
def minmax(L):
    assert hasattr(L, "__iter__")
    #
    # YOUR CODE HERE
    #

```

```python
# `minmax_test`: Test cell

L = [8, 7, 2, 5, 1]
mmL = minmax(L)
mmL_true = (1, 8)
print("minmax({}) -> {} [True: {}]".format(L, mmL, mmL_true))
assert type(mmL) is tuple and mmL == (1, 8)

from random import sample
L = sample(range(1000), 10)
mmL = minmax(L)
L_s = sorted(L)
mmL_true = (L_s[0], L_s[-1])
print("minmax({}) -> {} [True: {}]".format(L, mmL, mmL_true))
assert mmL == mmL_true

print("\n(Passed!)")
```

### Using Python to Work with Data Tables

Consider the following dataset of exam grades, organized as a 2-D table and stored in Python as a "list of lists" under the variable name, `grades`.

```python
grades = [
    # First line is descriptive header. Subsequent lines hold data
    ['Student', 'Exam 1', 'Exam 2', 'Exam 3'],
    ['Thorny', '100', '90', '80'],
    ['Mac', '88', '99', '111'],
    ['Farva', '45', '56', '67'],
    ['Rabbit', '59', '61', '67'],
    ['Ursula', '73', '79', '83'],
    ['Foster', '89', '97', '101']
]
import numpy as np

Newgrade = []
a = np.array(grades)
Newgrade.append(list(a[0:,1]))
Newgrade.append(list(a[0:,2]))
Newgrade.append(list(a[0:,3]))
```

#### Below we examine the contents of array ```a```

```python
a
```

#### Below we examine the contents of the ```Newgrade``` list

```python
Newgrade
```

**Exercise 5** (`students_test`). Write some code that computes a new list named `students[:]`, which holds the names of the students as they from "top to bottom" in the table.

```python
#
# YOUR CODE HERE
#

```

```python
# `students_test`: Test cell
print(students)
assert type(students) is list
assert students == ['Thorny', 'Mac', 'Farva', 'Rabbit', 'Ursula', 'Foster']
print("\n(Passed!)")
```

**Exercise 6** (`assignments_test`). Write some code to compute a new list named `assignments[:]`, to hold the names of the class assignments. (These appear in the descriptive header element of `grades`.)

```python
#
# YOUR CODE HERE
#

```

```python
# `assignments_test`: Test cell
print(assignments)
assert type(assignments) is list
assert assignments == ['Exam 1', 'Exam 2', 'Exam 3']
print("\n(Passed!)")
```

**BONUS Exercise {OPTIONAL}** (`avg_grades_by_student_test`). 

This exercise gives you practice working with Python _**Dictionaries**_. It is **NOT** required.

Write some code to create a dictionary named `avg_grades_by_student` that maps each student to his or her average exam score. For instance, `avg_grades_by_student['Thorny'] == 90`.

> **Hint.** The [`statistics`](https://docs.python.org/3.5/library/statistics.html) module of Python has at least one helpful function.

```python
# Create a dict mapping names to grade averages.
#
# YOUR CODE HERE
#

```

```python
# `avg_grades_by_student_test`: Test cell
print(avg_grades_by_student)
assert type(avg_grades_by_student) is dict, "Did not create a dictionary."
assert len(avg_grades_by_student) == len(students), "Output has the wrong number of students."
assert abs(avg_grades_by_student['Mac'] - 99.33333333333333) <= 4e-15, 'Mean is incorrect'
assert abs(avg_grades_by_student['Foster'] - 95.66666666666667) <= 4e-15, 'Mean is incorrect'
assert abs(avg_grades_by_student['Farva'] - 56) <= 4e-15, 'Mean is incorrect'
assert abs(avg_grades_by_student['Rabbit'] - 62.333333333333336) <= 4e-15, 'Mean is incorrect'
assert abs(avg_grades_by_student['Thorny'] - 90) <= 4e-15, 'Mean is incorrect'
assert abs(avg_grades_by_student['Ursula'] - 78.33333333333333) <= 4e-15, 'Mean is incorrect'
print("\n(Passed!)")
```

#### READY TO SUBMIT?
You've reached the end of this notebook. Be sure to restart and run all cells again to **make sure all cells are working** when they run in order. Then submit your **completed** HTML to the submission  folder for this activity.